---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: VideoExercise
key: 395ee4919f
xp: 50
```

`@projector_key`
f856cb77c6a12d11051d52d4d28a9e5e

---

## Fitting a line to a binary response

```yaml
type: VideoExercise
key: a2781695f6
xp: 50
```

`@projector_key`
e2eaa8253f6beb165d878ed1a596c863
